# Exact Inference in Bayesian Networks

This notebook develops the mathematical foundations of probabilistic queries and exact inference in Bayesian networks.

## 1. Theory Exercises

### 1.1. Probabilistic Queries: Shapes and Counts

Consider a full joint distribution over four variables $(A, B, C, D)$ with cardinalities $|A|=2$, $|B|=3$, $|C|=5$, and $|D|=5$.

For each query below:

1. **Provide the shape** of the answer to the query.
   - Express the shape as a tuple of dimension sizes, in the same order as the variables in the query. Omit singleton dimensions.
   - Example: if the query is $P(D, A=1)$, the expected shape is `(5,)`.
2. **Provide the total number of non-redundant values** in the result.
3. **Provide the number of probability distributions** in the result (the number of complete conditional distributions that are “stacked together”). 

Queries:
1. $P(A \mid B)$  
2. $P(C, D \mid A=1, B=0)$  
3. $P(B \mid C=4)$  
4. $P(A=1 \mid D=4)$

<!-- grade_id: cell-e2fe20ddc052f38b, points: 0, grade: false, solution: true -->
```python
result = {
    "P(A|B)": {"shape": (2,3), "n_values": 6, "n_distributions": 3},
    "P(C,D|A=1,B=0)": {"shape": (5, 5), "n_values": 25, "n_distributions": 1},
    "P(B|C=4)": {"shape": (3,), "n_values": 3, "n_distributions": 1},
    "P(A=1|D=4)": {"shape": (), "n_values": 1, "n_distributions": 1},
}
```

<!-- grade_id: cell-af5e81d120bede69, points: 4, grade: true, solution: false -->
```python
def pt(k, f):
    if k not in result: raise AssertionError(f"Missing {k}")
    if f not in result[k]: raise AssertionError(f"{k}:{f} missing")
    v = result[k][f]
    if v is None: raise AssertionError(f"{k}:{f} is None")
```

<!-- grade_id: cell-1221fcb6e5bbd419, points: 1, grade: true, solution: false -->
```python
# format check
for f in ("shape","n_values","n_distributions"): pt("P(A|B)", f)
```

<!-- grade_id: cell-e1207e42a835a66a, points: 1, grade: true, solution: false -->
```python
for f in ("shape","n_values","n_distributions"): pt("P(C,D|A=1,B=0)", f)
```

<!-- grade_id: cell-03600a70d9ed9567, points: 1, grade: true, solution: false -->
```python
for f in ("shape","n_values","n_distributions"): pt("P(B|C=4)", f)
```

<!-- grade_id: cell-d10059db2ff39179, points: 1, grade: true, solution: false -->
```python
for f in ("shape","n_values","n_distributions"): pt("P(A=1|D=4)", f)
```

### 1.2. Counting Arithmetic Operations: Inference via Enumeration vs. Variable Elimination

We continue with the same FJDT over $(A,B,C,D)$ (cardinalities $|A|=2$, $|B|=3$, $|C|=5$, $|D|=5$) which factorizes into the following conditional probability distributions:
- $P(A)$
- $P(B \mid A)$
- $P(C \mid A)$
- $P(D \mid B C)$

For each query below, report the **minimum** number of multiplication operations needed for inference via **Enumeration** and **Variable Elimination**.
- `n_op_enum`: the total number of **additions and multiplications** for a naive implementaiton of inference by enumeration which constructs atomic probabilites from the conditional probability distribution tables and sum over the hidden variables.
- `n_op_elim`: the total number of **additions and multiplications** for inference by variable elimination.

Queries:
1. $P(A \mid B=\lnot b)$  
2. $P(A=1 \mid D=4,\; C=0)$  
3. $P(A,\; D \mid B,\; C)$

<!-- grade_id: cell-dad5d9129b57b696, points: 0, grade: false, solution: true -->
```python
op_min = {
    "P(A|B=not b)":      {"n_op_enum": 225, "n_op_elim": 125},
    "P(A=1|D=4,C=0)":    {"n_op_enum": 24,  "n_op_elim": 16},
    "P(A,D|B,C)":        {"n_op_enum": 471,  "n_op_elim": 471},
}
```

<!-- grade_id: cell-0d0b1339fc5e36da, points: 0, grade: true, solution: false -->
```python
# shared helpers
def pt(k, f):
    if k not in op_min: raise AssertionError(f"Missing {k}")
    if f not in op_min[k]: raise AssertionError(f"{k}:{f} missing")
    v = op_min[k][f]
    if v is None: raise AssertionError(f"{k}:{f} is None")
    if type(v) is not int: raise AssertionError(f"{k}:{f} must be int")
```

<!-- grade_id: cell-a722b4e2609f2624, points: 0.5, grade: true, solution: false -->
```python
pt("P(A|B=not b)", "n_op_enum")
```

<!-- grade_id: cell-d999af7fd080093a, points: 0.5, grade: true, solution: false -->
```python
pt("P(A|B=not b)", "n_op_elim")
```

<!-- grade_id: cell-2a396d37da1c416e, points: 0.5, grade: true, solution: false -->
```python
pt("P(A=1|D=4,C=0)", "n_op_enum")
```

<!-- grade_id: cell-1a02d001723005d4, points: 0.5, grade: true, solution: false -->
```python
pt("P(A=1|D=4,C=0)", "n_op_elim")
```

<!-- grade_id: cell-8835a8f68b6e02e6, points: 0.5, grade: true, solution: false -->
```python
pt("P(A,D|B,C)", "n_op_enum")
```

<!-- grade_id: cell-d8f9aa5e2bd77a43, points: 0.5, grade: true, solution: false -->
```python
pt("P(A,D|B,C)", "n_op_elim")
```

### 1.3. (Conditional) Independence in Bayesian Networks

Consider the following Bayesian network:

<img src="bn.svg" style="height:430px" />

Store `True` into the corresponding result variable if the answer must be True and `False` otherwise.

) if you cannot intuitively solve these problems.* 

 1. $I \perp I$ ?
 2. $L \perp C$ ?
 3. $D \perp I$ ?
 4. $L \perp A \mid D$ ?
 5. $K \perp C \mid I, B$ ?
 6. $G \perp O \mid R$ ?

<!-- grade_id: cell-6fc4bbe72b927758, points: 0, grade: false, solution: true -->
```python
 # Provide either True or False as an answer
results = {}
results['1.3.1'] = False
results['1.3.2'] = False
results['1.3.3'] = False
results['1.3.4'] = False
results['1.3.5'] = True
results['1.3.6'] = False
```

<!-- grade_id: cell-6de53538ddbb1ce9, points: 0, grade: true, solution: false -->
```python
assert results.get("1.3.1") is not None
```

<!-- grade_id: cell-affd088758dbb225, points: 0.5, grade: true, solution: false -->
```python
assert results.get("1.3.2") is not None
```

<!-- grade_id: cell-d9d0f7b49d9d190d, points: 0.5, grade: true, solution: false -->
```python
assert results.get("1.3.3") is not None
```

<!-- grade_id: cell-d165cf4ef6f9d4d8, points: 0.5, grade: true, solution: false -->
```python
assert results.get("1.3.4") is not None
```

<!-- grade_id: cell-360ff7d6851dc526, points: 0.5, grade: true, solution: false -->
```python
assert results.get("1.3.5") is not None
```

<!-- grade_id: cell-dbfb6b8c9be56192, points: 0.5, grade: true, solution: false -->
```python
assert results.get("1.3.6") is not None
```